# CS 3990/5990: Secure Distributed Computation
## In-Class Exercise, Week of 9/11/2023

```python
# Imports and definitions
import numpy as np
from collections import defaultdict
import galois

GF = galois.GF(97)
```

# Questions from Last Week

## Question 4

Consider a semi-honest adversary. In the real world, what information is included in the view of the (potentially corrupted) party $p_i$?

<!-- grade_id: cell-813b2fb7ee8bdac3, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 5

Describe a *simulator* for the ideal world, that constructs views which are indistinguishable from the ones described in Question 1.

<!-- grade_id: cell-368c746db8f210b1, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

```python
from collections import defaultdict
class Party:
    """A participant in a multiparty computation protocol."""
    def __init__(self):
        """Initialize the dictionary to hold received messages."""
        self.received = defaultdict(list)
    
    def send(self, other, round, msg):
        """Simulate sending a message `msg` to another party `other` during round `round`"""
        other.received[round].append(msg)
```

## Question 6

Implement the `InsecureAggregationParty` class, which adds up the parties' inputs *insecurely*.

<!-- grade_id: cell-253770a6a53e06e8, points: 0, grade: false, solution: true -->
```python
class InsecureAggregationParty(Party):
    def round1(self, parties, input):
        self.input = GF(input)
        
        # YOUR CODE HERE
        raise NotImplementedError()

    def round2(self):
        # YOUR CODE HERE
        raise NotImplementedError()

    def get_view(self):
        return (self.input, self.output, self.received[1])
```

<!-- grade_id: cell-41c0a04587798b8c, points: 5, grade: true, solution: false -->
```python
# TEST CASE for question 1

# 5 parties
parties = [InsecureAggregationParty() for _ in range(5)]

# run round 1
for party in parties:
    party.round1(parties, 5)

# run round 2 and output
for party in parties:
    party.round2()
    #print('Party output:', party.output, type(party.output))
    assert party.output == GF(25)
```

## Question 7

Implement a simulator for the `InsecureAggregationParty` protocol, which constructs the view for party $i$ indistinguishable from the real-world view of party $i$ in the protocol, using only the inputs and output of the ideal functionality.

<!-- grade_id: cell-a9f12612df1fdec4, points: 0, grade: false, solution: true -->
```python
def simulator(n, i, inputs, output):
    """Simulates a real-world view in the ideal world. Outputs a 
    3-tuple: (input, output, received messages from round 1)"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1feb084cc713da72, points: 5, grade: true, solution: false -->
```python
# TEST CASE for question 2

# the real-world view of each party should be indistinguishable
# from the simulated view in the ideal world
# there is no randomness here, so they are really equal!
all_inputs = [5 for _ in range(5)]
for i, party in enumerate(parties):
    assert party.get_view() == simulator(5, i, all_inputs, 25)
```

## Question 8

Can you construct a simulator for `InsecureAggregationParty` without using the input for all $n$ parties? If not, why not? What does this mean about the security of `InsecureAggregationParty`?

<!-- grade_id: cell-be4f80bbafa34239, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

# New Questions

## Question 1

Write code to generate shares of a secret $x$ in a $(t, n)$-secret sharing scheme using Shamir's technique, where $n = 5$ and $t = 2$.

<!-- grade_id: cell-5009ee71e62855fc, points: 0, grade: false, solution: true -->
```python
def shr_2_5(v):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Example for question 1

shr_2_5(GF(5))
```

<!-- grade_id: cell-555fb4fef5fb1a22, points: 5, grade: true, solution: false -->
```python
# TEST CASE for question 1

assert len(shr_2_5(GF(5))) == 5
```

## Question 2

Write a function to reconstruct the secret, using only two shares.

<!-- grade_id: cell-e32dd909f78b8f42, points: 0, grade: false, solution: true -->
```python
def reconstruct(s1, s2):
    # YOUR CODE HERE
    raise NotImplementedError()

shares = shr_2_5(GF(5))
reconstruct(shares[0], shares[1])
```

<!-- grade_id: cell-65162c80c83feaa7, points: 5, grade: true, solution: false -->
```python
# TEST CASE
shares = shr_2_5(GF(5))
assert reconstruct(shares[0], shares[1]) == GF(5)
```

## Question 3

Why is a threshold secret sharing scheme more useful than the simpler additive secret sharing scheme we saw earlier?

<!-- grade_id: cell-021796358da78ff7, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 4

Write code to generate shares of a secret $x$ in a $(t, n)$-secret sharing scheme using Shamir's technique, for any $t$ and $n$.

<!-- grade_id: cell-34c3029e58f042c0, points: 0, grade: false, solution: true -->
```python
def shamir_share(v, t, n):
    # YOUR CODE HERE
    raise NotImplementedError()

shamir_share(GF(5), 3, 6)
```

```python
# Example for question 1

shr_2_5(GF(5))
```

<!-- grade_id: cell-a6374001a42a98cf, points: 5, grade: true, solution: false -->
```python
# TEST CASE
assert len(shamir_share(GF(5), 3, 6)) == 6
shares = shamir_share(GF(5), 2, 6)
assert reconstruct(shares[0], shares[1]) == GF(5)
```

## Question 5

Given the two sets of shares `shares1` and `shares2` below, write a function whose output is their sum (as a set of shares).

<!-- grade_id: cell-96a3bb1415b78611, points: 0, grade: false, solution: true -->
```python
shares1 = shamir_share(GF(20), 2, 6)
shares2 = shamir_share(GF(5), 2, 6)

def add_shares(shares1, shares2):
    # YOUR CODE HERE
    raise NotImplementedError()

added_shares = add_shares(shares1, shares2)
print(added_shares)
reconstruct(added_shares[0], added_shares[1])
```

<!-- grade_id: cell-fc07ee60a24605f8, points: 5, grade: true, solution: false -->
```python
# TEST CASE
added_shares = add_shares(shares1, shares2)
assert reconstruct(added_shares[0], added_shares[2]) == GF(25)
```

## Question 6

Write a function to reconstruct a secret from a set of at least $t$ shares. Use the `galois.lagrange_poly` function, which implements [Lagrange interpolation](https://en.wikipedia.org/wiki/Lagrange_polynomial).

<!-- grade_id: cell-8e6f22549d0766db, points: 0, grade: false, solution: true -->
```python
def reconstruct(shares):
    # YOUR CODE HERE
    raise NotImplementedError()

reconstruct(added_shares)
```

<!-- grade_id: cell-0fc28d5b0ecdae18, points: 5, grade: true, solution: false -->
```python
# TEST CASE
shares = shamir_share(GF(30), 5, 10)
assert reconstruct(shares) == GF(30)
assert reconstruct(shares[:5]) == GF(30)  # t shares are sufficient
assert reconstruct(shares[:4]) != GF(30)  # t - 1 shares are not sufficient
```

## Question 7

Given the two sets of shares `shares1` and `shares2` below, write a function whose output is their product (as a set of shares).

<!-- grade_id: cell-27b6a75239e0abfc, points: 0, grade: false, solution: true -->
```python
shares1 = shamir_share(GF(20), 3, 6)
shares2 = shamir_share(GF(3), 3, 6)

def mult_shares(shares1, shares2):
    # YOUR CODE HERE
    raise NotImplementedError()

product_shares = mult_shares(shares1, shares2)
print(product_shares)
reconstruct(shares1)
reconstruct(product_shares)
```

<!-- grade_id: cell-a208fe44dcacec08, points: 5, grade: true, solution: false -->
```python
# TEST CASE
product_shares = mult_shares(shares1, shares2)

assert reconstruct(product_shares) == GF(60)
assert reconstruct(product_shares[:4]) != GF(60)  # t shares are no longer sufficient
```